# Analytics DL 테스트

```python
import numpy as np
import tensorflow as tf
```

## Dense Layer 만들기

`input_x`와 `output_dim`을 인자로 받는 Dense Layer Function을 만들어 보자.  
`bias`가 있어야 한다.


<div class="alert alert-block alert-info"  markdown="1">
<b>$\divideontimes$ 설명 </b> 
<br>

* Argument
  - `input_x` : `numpy.ndarray`
  - `output_dim` : `int`
  - `seed` : `1`


* Return
  - `output`: `numpy.ndarray`
  - `weight` : `numpy.ndarray`
  - `bias` : `numpy.ndarray`


</div>

<!-- grade_id: 01_dense_code, points: 0, grade: false, solution: true -->
```python
def dense_layer(
    input_x,
    output_dim=None,
    seed=1,
    ):
    input_dim = input_x.shape[-1]
    np.random.seed(seed)
    
    """
    Hint: np.random.random((dim, dim)) 을 사용하자.
    
    """

    ### BEGIN SOLUTION

    weight = np.random.random((input_dim, output_dim))
    bias = np.random.random((output_dim, ))

    output = input_x @ weight + bias
     
    ### END SOLUTION

    return output, weight, bias
```

```python
"""테스트용 셀"""
```

<!-- grade_id: 01_dense_test, points: 10, grade: true, solution: false -->
```python
"""채점용 셀(삭제금지)"""
dense_layer
### BEGIN HIDDEN TESTS
input_a = np.random.random((2, 3)).astype(np.float32)
input_b = np.random.random((7, 5)).astype(np.float32)

oa, wa, ba = dense_layer(input_a, output_dim=5)
ob, wb, bb = dense_layer(input_b, output_dim=2)
assert np.array_equal(oa, input_a @ wa + ba)
assert np.array_equal(ob, input_b @ wb + bb)
### END HIDDEN TESTS
```

## Activation Function

`activation function`은 딥러닝 모델에 비선형성을 부여하고,  
값을 일정 범위 내로 bound시켜 학습을 용이하게 하는 기능도 있다.

많은 `activation function` 중  
`sigmoid`와 `tanh` 의 값의 범위를 각각 복습해 보자.

<!-- grade_id: 02_activation_code, points: 0, grade: false, solution: true -->
```python
def q_activation():
    """
    Return
      sigmoid_range = (min, max)
      tanh_range = (min, max)
    
    """
    ### BEGIN SOLUTION
    
    sigmoid_range = (0, 1)
    tanh_range = (-1, 1)

    ### END SOLUTION
    
    return sigmoid_range, tanh_range
```

```python
"""테스트용 셀"""
```

<!-- grade_id: 02_activation_test, points: 10, grade: true, solution: false -->
```python
"""채점용 셀(삭제금지)"""
q_activation
### BEGIN HIDDEN TESTS
sigmoid_range, tanh_range = q_activation()
assert all(map(lambda x: len(x) == 2, [sigmoid_range, tanh_range]))
assert (min(sigmoid_range) == 0) and (max(sigmoid_range) == 1)
assert (min(tanh_range) == -1) and (max(tanh_range) == 1)
### END HIDDEN TESTS
```

## Tensorflow Session

`tensorflow`는 Graph와 Flow의 개념을 합쳐 만들어진 Framework이다.  
Flow를 실행하기 위한 `tf.Session` 사용법을 다시 확인해 보자.

<!-- grade_id: 03_tfsession_code, points: 0, grade: false, solution: true -->
```python
np.random.seed(1)
arr_a = np.random.random((4, 3)).astype(np.float32)
arr_b = np.random.random((3, 5)).astype(np.float32)


def q_tfsession(a, b):
    """
    a, b 행렬을 받아서 행렬곱을 수행하는 함수를 Tensorflow로 만들어 보자.
    """
    tf.reset_default_graph()
    
    tf_a = tf.constant(a, dtype=tf.float32)
    tf_b = tf.constant(b, dtype=tf.float32)
    
    matmal_op = tf_a @ tf_b  # tf.matmul(tf_a, tf_b)

    ### BEGIN SOLUTION
    
    with tf.Session() as sess:
        
        answer = sess.run(matmal_op)
    
    ### END SOLUTION

    return answer
```

```python
"""테스트용 셀"""
```

<!-- grade_id: 03_tfsession_test, points: 10, grade: true, solution: false -->
```python
"""채점용 셀(삭제금지)"""
#q_tfsession
### BEGIN HIDDEN TESTS
for _ in range(5):

    np.random.seed(1)
    test_arr_a = np.random.random((_ + 17, _ + 13)).astype(np.float32)
    test_arr_b = np.random.random((_ + 13, _ + 29)).astype(np.float32)

    test_answer = q_tfsession(test_arr_a, test_arr_b)
    assert test_answer.shape == (test_arr_a @ test_arr_b).shape
### END HIDDEN TESTS
```

## Tensorflow Graph & Batch

Tensorflow Graph에서는 외부 변수를 받을 수 있도록 `tf.placeholder`를 사용한다.  
또한 학습을 위해 `batch`라는 개념을 활용하여 데이터를 조금씩 학습해 나간다.

`tf.placeholder`와 `batch`를 활용하는 함수를 작성해 보자.
편의상 input 전체 사이즈는 짝수로 한정한다.


<!-- grade_id: 04_tfgraphbatch_code, points: 0, grade: false, solution: true -->
```python
np.random.seed(1)
input_a = np.random.random((8, 3)).astype(np.float32)
input_b = np.random.random((8, 3)).astype(np.float32)

def q_tfgraph_and_batch(a, b, batch_size=2):
    
    tf.reset_default_graph()
    
    batch_num = int(a.shape[0] / batch_size)
    a_for_batch = np.split(a, batch_num)
    b_for_batch = np.split(b, batch_num)
    
    """
    Hint : tf.placeholder(dtype=tf.float32, shape=(None, ))

    """
    ### BEGIN SOLUTION
    
    tf_a = tf.placeholder(tf.float32, shape=(None, a.shape[1]))
    tf_b = tf.placeholder(tf.float32, shape=(None, b.shape[1]))

    ### END SOLUTION
    
    add_op = tf_a + tf_b  # tf.matmul(tf_a, tf_b)
    
    
    """
    Hint : Batch로 나온 결과를 answer_list에 담자.

    """
    answer_list = []
    ### BEGIN SOLUTION
    with tf.Session() as sess:
        
        for batch_a, batch_b in zip(a_for_batch, b_for_batch):
            answer = sess.run(add_op, feed_dict={tf_a: batch_a, tf_b: batch_b})
            answer_list += [answer]

    ### END SOLUTION
    
    concat_answer = np.concatenate(answer_list)

    return concat_answer
```

```python
"""테스트용 셀"""
```

<!-- grade_id: 04_tfgraphbatch_test, points: 10, grade: true, solution: false -->
```python
"""채점용 셀(삭제금지)"""
q_tfgraph_and_batch
### BEGIN HIDDEN TESTS
for _ in range(1, 6):

    np.random.seed(1)
    test_arr_a = np.random.random((_ * 10, _ + 13)).astype(np.float32)
    test_arr_b = np.random.random((_ * 10, _ + 13)).astype(np.float32)

    test_answer = q_tfgraph_and_batch(test_arr_a, test_arr_b)
    assert test_answer[_].shape == (test_arr_a[_] + test_arr_b[_]).shape
### END HIDDEN TESTS
```

## Tensorflow RNN

`Bidirectional RNN`은 `RNN`의 발전된 형태이다.  
이 구조를 기반으로 `RNN`을 역으로 구현해 보자.

<!-- grade_id: 05_tfrnn_code, points: 0, grade: false, solution: true -->
```python
def rnn_fn(
    num_units,
    input_x,
    output_dim,
    
    ):
    
    tf.reset_default_graph()
    """BIRNN EXAMPLE
    
    1. Cell 생성
    rnn_forward_cell = tf.nn.rnn_cell.BasicRNNCell(
        num_units,
    )
    rnn_backward_cell = tf.nn.rnn_cell.BasicRNNCell(
        num_units,
    )
    
    
    2. biRNN Flow 생성
    (outputs_fw, outputs_bw), _state = tf.nn.bidirectional_dynamic_rnn(
        cell_fw=rnn_forward_cell,
        cell_bw=rnn_backward_cell,
        inputs=input_x,
        dtype=tf.float32,
    )
    rnn_outputs = outputs_fw + outputs_bw
    
    
    3. 예측값 생성
    predict = tf.layers.dense(
        rnn_outputs,
        units=output_dim,
        activation=tf.nn.tanh,
    )
    
    """
    
    # 1. RNN Cell 생성
    
    ### BEGIN SOLUTION
    rnn_cell = tf.nn.rnn_cell.BasicRNNCell(
        num_units,
    )
    ### END SOLUTION

    
    # 2. Bidirectional RNN Flow 생성
    
    ### BEGIN SOLUTION
    rnn_outputs, _state = tf.nn.dynamic_rnn(
        rnn_cell,
        inputs=input_x,
        dtype=tf.float32,
    )
    ### END SOLUTION


    # 3. 예측값 생성 
    predict_y = tf.layers.dense(
        inputs=rnn_outputs[:, -2:, :],
        units=output_dim,
        activation=tf.nn.tanh,
    )
    
    
    # return 예측값 
    return predict_y, rnn_outputs, _state
```

```python
"""테스트용 셀"""
```

<!-- grade_id: 05_tfrnn_test, points: 10, grade: true, solution: false -->
```python
"""채점용 셀(삭제금지)"""
rnn_fn
### BEGIN HIDDEN TESTS


for _ in range(1, 6):

    np.random.seed(1)
    x = np.array(np.random.random((12 * _, 5, 3)), dtype=np.float32)
    y = np.array(np.random.random((12 * _, 2, 2)), dtype=np.float32)
    rnn_y = np.array(np.random.random((12 * _, 5, 4)), dtype=np.float32)

    test_predict_y, test_rnn_outputs, test_state = rnn_fn(
        num_units=4,
        input_x=x,
        output_dim=2,
    )

    assert test_predict_y.shape == (12 * _, 2, 2)
    assert test_rnn_outputs.shape == (12 * _, 5, 4)

### END HIDDEN TESTS
```

End of Document.